<a href="https://colab.research.google.com/github/tsangrebecca/BloomTech/blob/main/Sprint_1/Module_1/Intro_Python_Assignment_AG_completed.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Unit 1 Sprint 1 Module 1

## Introduction to Python

### Objectives

* Create and work with variable data types
* Understand indexing and methods
* Practice dictionary indexing

#### Total notebook points: 8

## Autograded Module Projects

Welcome to the first Module Project of Unit 1! You will complete a project (sometimes also referred to as an assignment) after the Guided Project for each Module. There will be four Module Projects per Sprint and each project is designed to provide you with an opportunity to practice what you learned in the Canvas Warm-up material and the Guided Project with your instructors.

Throughout Unit 1, the Module Projects and the Sprint Challenges are *autograded*. You will complete your work in a Jupyter/Python notebook (the files that end with `.ipynb`) and then upload your completed notebook to Canvas and submit for grading. This autograding process will check your answers and provide more information about the errors in your notebook. You'll receive a score when the testing is complete.

So, let's get started! If you are reading this notebook, then you have the correct autograded version.

**Task 1** - Write Your Name

Create a variable called `your_name` that includes your name. 

<!-- grade_id: cell-6027a63ecafcefab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
your_name = "rebecca tsang"
```

**Task 1 Test**

<!-- grade_id: cell-a6c4f852066437b1, points: 1, grade: true, solution: false -->
```python
#Task 1 - Test
assert isinstance(your_name, str)
print("Hello", your_name )



```

Output:
```
Hello rebecca tsang
```

**Task 2** - Variable data types

* Assign the number `10` to a variable called `first_num`
* Assign the number `5` to a variable called `second_num`
* Assign the number `2` to a variable called `num_two` 

<!-- grade_id: cell-e96ddb74e7505094, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
first_num = 10
second_num = 5
num_two = 2
```

**Task 2 Test**

<!-- grade_id: cell-a33f13464286f445, points: 1, grade: true, solution: false -->
```python
#Task 2 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 3** - Operators
* Return the product of all three variables created in the task above. Assign it to an object called `product_total`.
* Return the floor division of `product_total` by 7. Assign it **as an interger** to an object called `floor_value`. 

<!-- grade_id: cell-311e511f461035fc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
product_total = first_num * second_num * num_two
product_total
```

Output:
```
100
```

```python
floor_value = product_total//7
floor_value
```

Output:
```
14
```

**Task 3 Test**

<!-- grade_id: cell-c96d884a0d9c28fe, points: 1, grade: true, solution: false -->
```python
#Task 3 Test

assert product_total == 100, "Check your operators."

# Hidden tests - you will see the results when you submit to Canvas

```

**Task 4a** Python Data Types
* Create a list of five objects: `red`, `green`, `blue`, `yellow`, `purple`. Assign this list **in alphabetical order** to an object called `five_colors`.  


<!-- grade_id: cell-32200b8130c2bd81, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
five_colors = ["red", "green", "blue", "yellow", "purple"]
five_colors.sort()
five_colors
```

Output:
```
['blue', 'green', 'purple', 'red', 'yellow']
```

**Task 4a Test**

<!-- grade_id: cell-041962ba22262d5e, points: 0.5, grade: true, solution: false -->
```python
#Task 4a Test
assert len(five_colors) == 5
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 4b** Python Data Types
* Remove the value that is in the middle of the `five_colors` list. Assign this new list to an object called `four_colors`.

<!-- grade_id: cell-de568d701f9137ca, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
removed_color = five_colors.pop(2)
print(removed_color)
```

Output:
```
purple
```

```python
four_colors = five_colors
four_colors
```

Output:
```
['blue', 'green', 'red', 'yellow']
```

**Task 4b Test**

<!-- grade_id: cell-9ab8dc3c20c2a9c1, points: 0.5, grade: true, solution: false -->
```python
#Task 4b Test
assert len(four_colors) == 4, "Try displaying or printing the output of your object to verify your results"
# Hidden tests - you will see the results when you submit to Canvas

```

**Task 5** - Indexing
* Use indexing to return the *third* item in the `four_colors` list. Assign this to an object called `third_color`.
* Next, use string indexing to return the first and the last letter of the `third_color` object. Assign it to an object called `first_and_last`.   


```python
third_color = four_colors[2]
print(third_color)
```

Output:
```
red
```

```python
first_and_last = third_color[0] + third_color[-1]
print(first_and_last)
```

Output:
```
rd
```

**Task 5 Test**

<!-- grade_id: cell-229ed9f4a0babeba, points: 1, grade: true, solution: false -->
```python
#Task 5 Test
assert third_color == 'red'
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 6a** - Indexing
* Create a dictionary that includes the following key: value pairs: `charizard: fire, squirtle: water, bulbasaur: leaf`. Assign it to an object called `poke_dictionary` 

<!-- grade_id: cell-24b1a6230e320ec9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
poke_dictionary = {"charizard": "fire", "squirtle": "water", "bulbasau": "leaf"}
```

**Task 6a Test**

<!-- grade_id: cell-5ad0aaf04fb6881b, points: 0.5, grade: true, solution: false -->
```python
#Task 6a Test
assert len(poke_dictionary.keys()) == 3, "Verify your syntax. Did you create a dictionary object?"

# Hidden tests - you will see the results when you submit to Canvas

```

**Task 6b** - Indexing
* Create a new dictionary called `more_poke` that includes the following: `pikachu: electric, pidgey: air` 
* Combine the key and values of the `more_poke` with the original `poke_dictionary`.
>Hint: You may need to research how to combine two dictionaries. 
* Assign all the `dictionary values` of the `poke_dictionary` to an object called `poke_values`. 
>Hint: Your output for `poke_values` will look like: `dict_values(['fire', 'water', 'leaf', 'electric', 'air'])`

<!-- grade_id: cell-0b4582e6a68091ef, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
more_poke = {"pikachu": "electric", "pidgey": "air"}
```

```python
poke_dictionary.update(more_poke)
```

```python
poke_values = poke_dictionary.values()
print(poke_values)
```

Output:
```
dict_values(['fire', 'water', 'leaf', 'electric', 'air'])
```

**Task 6b Tests**

<!-- grade_id: cell-c561f60e90446fd7, points: 0.5, grade: true, solution: false -->
```python
#Task 6b Test
assert isinstance(more_poke, dict), "Check your syntax. Did you create a dictionary object?"
assert len(poke_dictionary) == 5, "Did you properly combine both dictionaries?"
# Hidden tests - you will see the results when you submit to Canvas



```

**Task 7** Loops
* Use a for loop to go through the `poke_dictionary` object. Return all the dictionary **values** as uppercase letters and the **keys** as lowercase values. 
* *Hint*: This [StackOverflow](https://stackoverflow.com/questions/50004310/convert-python-dictionary-to-uppercase) page offers a solution (with a slight adjustment), but there are multiple ways you can do this!

```python
poke_dictionary
```

Output:
```
{'charizard': 'fire',
 'squirtle': 'water',
 'bulbasau': 'leaf',
 'pikachu': 'electric',
 'pidgey': 'air'}
```

```python
poke_dictionary.items()
```

Output:
```
dict_items([('charizard', 'fire'), ('squirtle', 'water'), ('bulbasau', 'leaf'), ('pikachu', 'electric'), ('pidgey', 'air')])
```

```python
# YOUR CODE HERE
# USE PRINT FUNCTION TO CHECK CODE IN FOR LOOPS, we don't need print statement inside the loop, but good to use it to check what's going on.

for i in poke_dictionary:
  poke_dictionary[i] = poke_dictionary[i].upper()      # for all the values in the poke_dictionary by looping each pokemon and returning their values, then reassign them.
  
```

<!-- grade_id: cell-0535b02af29e01e7, points: 0, grade: false, solution: true -->
```python
# another way to do it, if we want to change both keys AND values in dictionary
#new_dict = {}

#for key, value in poke_dictionary.items():   # use .item() method to iterate over a dictionary and use "key" and "value" to have access to the keys and values simultaneously.
#   key_lowercase = key.lower()              # create keys all lower case
#   value_uppercase = value.upper()          # create values all UPPER CASE
#   new_dict[key_lowercase] = value_uppercase   #put the keys and values back in the new dictionary
#print(new_dict)

# reassign new_dict to the poke_dictionary
# poke_dictionary = new_dict

```

```python
poke_dictionary
```

Output:
```
{'charizard': 'FIRE',
 'squirtle': 'WATER',
 'bulbasau': 'LEAF',
 'pikachu': 'ELECTRIC',
 'pidgey': 'AIR'}
```

**Task 7 Test**

<!-- grade_id: cell-0e4a39d38c38f186, points: 1, grade: true, solution: false -->
```python
#Task 8 Test
assert poke_dictionary['pikachu'].isupper(), "Your dictionary values should all be uppercase"
```

**Task 8** Imports
* Import the `seaborn` and `pandas` library. Alias the libraries as `sns` and `pd` respectively. 
* Use the `sns` library to load a dataset of your choice. Assign it to an object called `df`. 
* Run the following code: `df.loc[0]`. What data type is this? Assign your answer to an object called `first_row_type`.


<!-- grade_id: cell-6ed26cd55b737be1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import seaborn as sns
import pandas as pd
```

```python
df = sns.load_dataset('diamonds')
df
```

Output:
```
carat        cut color clarity  depth  table  price     x     y     z
0       0.23      Ideal     E     SI2   61.5   55.0    326  3.95  3.98  2.43
1       0.21    Premium     E     SI1   59.8   61.0    326  3.89  3.84  2.31
2       0.23       Good     E     VS1   56.9   65.0    327  4.05  4.07  2.31
3       0.29    Premium     I     VS2   62.4   58.0    334  4.20  4.23  2.63
4       0.31       Good     J     SI2   63.3   58.0    335  4.34  4.35  2.75
...      ...        ...   ...     ...    ...    ...    ...   ...   ...   ...
53935   0.72      Ideal     D     SI1   60.8   57.0   2757  5.75  5.76  3.50
53936   0.72       Good     D     SI1   63.1   55.0   2757  5.69  5.75  3.61
53937   0.70  Very Good     D     SI1   62.8   60.0   2757  5.66  5.68  3.56
53938   0.86    Premium     H     SI2   61.0   58.0   2757  6.15  6.12  3.74
53939   0.75      Ideal     D     SI2   62.2   55.0   2757  5.83  5.87  3.64

[53940 rows x 10 columns]
```

```python
df.loc[0]
```

Output:
```
carat       0.23
cut        Ideal
color          E
clarity      SI2
depth       61.5
table       55.0
price        326
x           3.95
y           3.98
z           2.43
Name: 0, dtype: object
```

```python
first_row_type = type(df.loc[0])   # use the built-in type(xyz) function
print(first_row_type)
```

Output:
```
<class 'pandas.core.series.Series'>
```

**Task 8 Test**

<!-- grade_id: cell-39cc971493c3004e, points: 1, grade: true, solution: false -->
```python
#Task 8 Test

assert isinstance(df, pd.DataFrame), "Did you import the pandas library? Did you check the datatype of the df object?"

# Hidden tests - you will see the results when you submit to Canvas
```

 A class is a blueprint or a template that defines the structure and behavior, while an object is a specific instance created based on the class definition, with its own unique attributes and ability to perform methods defined within the class. The row of data is an object within the series class.

```python

```